# nbgrader

nbgrader is a Jupyter notebook grading system. With nbgrader, you can add auto-graded or manually graded asessments. Manually graded assessments can be code-based or free response. You have the ability to set the number of points for each question, and you can override a grade if need be. Assessments are added directly to a notebook, and there is no need to do a system exit (all or nothing questions) or an HTTP request (partial points) as is required for advanced code tests in the Guide. Be sure to check out the [Codio documentation](https://docs.codio.com/courses/g
rading/#nbgrader-autograding) on nbgrader for more information.

## `.codio-jupyter` File

Codio requires a `.codio-jupyter`file in the file tree. Right-click on the root of the file tree and add this file. This file can be left empty, or you can set several [configuration options](https://docs.codio.com/courses/grading/#configuration).

## Creating an nbgrader Cell

Add a code cell to the notebook. In the menu bar click `View`, then `Cell Toolbar`, and finally click `Create Assignment`.

![nbgrader Cell](img/nbgrader_cell.png)

In the top-right corner of the code cell, you will see a dropdown menu. You can select the type of assessment you want.

![Cell Toolbar](img/cell_toolbar.png)
dsfasdf

asdfasdfasd asd asd
asdfasdf
asdfasdf
sadfasd fsadf asdf sadf
asdfasdfasdf

## Manually Graded Assessments

Select `Manually graded answer` from the dropdown menu in the cell toolbar and assign the number of points for the assessment. The ID for this cell is automatically generated, but you can change it to a more meaningful name. The cell below is an example of a manually graded assessment. Manually graded cells can be either markdown or coding cells.

<!-- grade_id: why_functions, points: 2, grade: true, solution: true -->
#  Why functions

Give two benefits of using functions in your code.

## Auto-Graded Assessments

Auto-graded assessments are divided into two types of cells, auto-graded answers and auto-graded tests. Auto-graded answer cells are cells for students to compose their answer. You can prepopulate these cells with boilerplate code if you would like. The following delimeters are used to designate an area for students to write their answer.

![Begin and End a Solution](img/begin-end-solution.png)

<!-- grade_id: cell-7abf9280766941df, points: 0, grade: false, solution: true -->
```python
# Complete the function below that returns a list of squares from 1 to n

def squares(n):
    """Calculates and returns a list of squares from 1 to n"""
    ### BEGIN SOLUTION
    answer = []
    for i in range(1, n+1):
        answer.append(i ** 2)
    return answer
    ### END SOLUTION
```

## Auto-Graded Tests

Grading student work is done with an auto-grader test cell. This cell applies several tests to student code using the `assert` command. Other tests `assert_equal` or `assert_raises`, but these need to be imported first.

![Testing Imports](img/test_imports.png)

Students will be able to see all of these test cases in the cell. This gives them an idea how their code will be tested. If you want to apply hidden tests to the student code, use the following delimeters.

![Hidden Test Syntax](img/hidden_solution.png)

<!-- grade_id: cell-90b9e6131983c587, points: 0, grade: true, solution: false -->
```python
"""Example of testing student code with visible and hidden tests"""
assert squares(1), [1]
assert squares(2), [1, 4]
assert squares(3), [1, 4, 9]
### BEGIN HIDDEN TESTS
assert squares(5), [1, 4, 9, 16, 25]
assert squares(10), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
### END HIDDEN TESTS
```

## Read-Only Cells

You can mark cells as read-only if you do not want students to alter them. Use the dropdown menu in the cell toolbar. Both markdown and coding cells can be marked as read-only. Testing cells will automatically be marked as read-only. Read-only cells will have a lock icon in the top-left corner. 

![Locked Cell](img/lock.png)

## Validating the Tests

The `Validate` button will verify that your tests are valid. If there is a problem, Jupyter will indicate the cell and line number.

![Test Failed](img/test_error.png)

## Partial Points

Auto-graded tests are all-or-nothing when it comes to assigning points. In the example above, there are a total of five tests and five points. However, if the student does not pass one of the tests, they are awarded 0 points. They must pass all of the tests to receive the 5 points. There are three ways to award partial points:

* One way around this is to break up each test into its own cell. Make each test cell worth 1 point. That way the student receives a point for each test passed.

* You could also have a manually graded cell for code. This assessment could be worth 5 points, and you would assign the grade you think the student earned.

* You can manually override a students grade when you open their submission.

## Feedback

Opening student submissions allows for you to write feedback. Below each problem (auto-graded or manually graded) is a text field where you can write your feedback.

![Student Feedback](img/feedback.png)